```python
import pandas as pd
import glob

all_files = glob.glob("./*.csv") 

```

```python
combined_df = pd.concat((pd.read_csv(file) for file in all_files), ignore_index=True)
```

Output:
```
/var/folders/s9/tjqjhlhs57q4zp00qz4178gc0000gn/T/ipykernel_28918/2762595841.py:1: DtypeWarning: Columns (13) have mixed types. Specify dtype option on import or set low_memory=False.
  combined_df = pd.concat((pd.read_csv(file) for file in all_files), ignore_index=True)
/var/folders/s9/tjqjhlhs57q4zp00qz4178gc0000gn/T/ipykernel_28918/2762595841.py:1: DtypeWarning: Columns (13) have mixed types. Specify dtype option on import or set low_memory=False.
  combined_df = pd.concat((pd.read_csv(file) for file in all_files), ignore_index=True)
/var/folders/s9/tjqjhlhs57q4zp00qz4178gc0000gn/T/ipykernel_28918/2762595841.py:1: DtypeWarning: Columns (13) have mixed types. Specify dtype option on import or set low_memory=False.
  combined_df = pd.concat((pd.read_csv(file) for file in all_files), ignore_index=True)
/var/folders/s9/tjqjhlhs57q4zp00qz4178gc0000gn/T/ipykernel_28918/2762595841.py:1: DtypeWarning: Columns (13) have mixed types. Specify dtype option on import or set low_memory=False.
  co
```

```python
combined_df = pd.concat((pd.read_csv(file, low_memory=False) for file in all_files), ignore_index=True)

```

```python
combined_df.to_csv('combined_file.csv', index=False)
```

```python
data = pd.read_csv('combined_file.csv')
data.head(15)
```

Output:
```
/var/folders/s9/tjqjhlhs57q4zp00qz4178gc0000gn/T/ipykernel_28918/3649526037.py:1: DtypeWarning: Columns (13,14) have mixed types. Specify dtype option on import or set low_memory=False.
  data = pd.read_csv('combined_file.csv')
    case_enquiry_id              open_dt        sla_target_dt  \
0      101003148265  2020-01-01 00:07:00  2020-01-13 03:30:00   
1      101003148266  2020-01-01 00:54:48  2020-01-01 00:54:48   
2      101003148268  2020-01-01 01:14:13  2020-01-03 03:30:00   
3      101003148269  2020-01-01 01:19:00  2020-01-03 03:30:00   
4      101003148271  2020-01-01 02:02:00  2020-01-03 03:30:00   
5      101003148272  2020-01-01 02:04:00  2020-01-03 03:30:00   
6      101003148273  2020-01-01 02:11:10  2020-01-03 03:30:00   
7      101003148274  2020-01-01 02:19:58                  NaN   
8      101003148276  2020-01-01 02:56:41  2020-01-03 03:30:00   
9      101003148277  2020-01-01 03:08:52  2020-01-06 03:30:00   
10     101003148278  2020-01-01 03:15:14  2020-01-03 03:3
```

```python
print("Column Names and Data Types:")
print(data.dtypes)
```

Output:
```
Column Names and Data Types:
case_enquiry_id                     int64
open_dt                            object
sla_target_dt                      object
closed_dt                          object
on_time                            object
case_status                        object
closure_reason                     object
case_title                         object
subject                            object
reason                             object
type                               object
queue                              object
department                         object
submitted_photo                    object
closed_photo                       object
location                           object
fire_district                      object
pwd_district                       object
city_council_district              object
police_district                    object
neighborhood                       object
neighborhood_services_district     object
ward                               object
preci
```

```python
# Check if there are rows where `submitted_photo` or `closed_photo` are not NaN
non_nan_rows = data[(~data['submitted_photo'].isna()) | (~data['closed_photo'].isna())]

# Display the result
print(non_nan_rows)

# Count the number of non-NaN rows
print(f"Number of rows where either column is not NaN: {len(non_nan_rows)}")

```

Output:
```
case_enquiry_id              open_dt        sla_target_dt  \
111         101003148392  2020-01-01 06:57:18  2020-03-09 04:30:00   
116         101003148397  2020-01-01 07:06:34  2020-03-09 04:30:00   
127         101003148409  2020-01-01 07:30:14  2020-03-09 04:30:00   
131         101003148413  2020-01-01 07:37:59  2020-03-09 04:30:00   
183         101003148471  2020-01-01 08:50:00  2020-01-16 03:30:00   
...                  ...                  ...                  ...   
2356200     101002279456  2017-12-05 21:08:31  2018-02-08 03:30:00   
2356202     101002279459  2017-12-05 21:09:56  2018-02-08 03:30:00   
2358566     101002282277  2017-12-09 20:11:00  2017-12-26 03:30:00   
2358576     101002282294  2017-12-09 23:22:00  2017-12-26 03:30:00   
2362691     101002287324  2017-12-16 21:07:35  2018-02-20 03:30:00   

                   closed_dt  on_time case_status  \
111      2020-09-27 04:14:02  OVERDUE      Closed   
116      2020-09-27 05:20:18  OVERDUE      Closed   
127      
```

```python
combined_df = pd.concat(
    (pd.read_csv(file, dtype={'submitted_photo': str, 'closed_photo': str}) for file in all_files),
    ignore_index=True
)
```

```python
combined_df['submitted_photo'] = combined_df['submitted_photo'].astype(str)
combined_df['closed_photo'] = combined_df['closed_photo'].astype(str)

```

```python
# Save the DataFrame to a new CSV file
combined_df.to_csv("combined_cleaned_data.csv", index=False)

```

```python
data1 = pd.read_csv('combined_cleaned_data.csv')
data1.head(15)
```

Output:
```
/var/folders/s9/tjqjhlhs57q4zp00qz4178gc0000gn/T/ipykernel_28918/1288416817.py:1: DtypeWarning: Columns (13,14) have mixed types. Specify dtype option on import or set low_memory=False.
  data1 = pd.read_csv('combined_cleaned_data.csv')
    case_enquiry_id              open_dt        sla_target_dt  \
0      101003148265  2020-01-01 00:07:00  2020-01-13 03:30:00   
1      101003148266  2020-01-01 00:54:48  2020-01-01 00:54:48   
2      101003148268  2020-01-01 01:14:13  2020-01-03 03:30:00   
3      101003148269  2020-01-01 01:19:00  2020-01-03 03:30:00   
4      101003148271  2020-01-01 02:02:00  2020-01-03 03:30:00   
5      101003148272  2020-01-01 02:04:00  2020-01-03 03:30:00   
6      101003148273  2020-01-01 02:11:10  2020-01-03 03:30:00   
7      101003148274  2020-01-01 02:19:58                  NaN   
8      101003148276  2020-01-01 02:56:41  2020-01-03 03:30:00   
9      101003148277  2020-01-01 03:08:52  2020-01-06 03:30:00   
10     101003148278  2020-01-01 03:15:14  2020-0
```

```python
data = pd.read_csv('combined_file.csv')
data.head(15)
```

Output:
```
/var/folders/s9/tjqjhlhs57q4zp00qz4178gc0000gn/T/ipykernel_28918/3649526037.py:1: DtypeWarning: Columns (13,14) have mixed types. Specify dtype option on import or set low_memory=False.
  data = pd.read_csv('combined_file.csv')
    case_enquiry_id              open_dt        sla_target_dt  \
0      101003148265  2020-01-01 00:07:00  2020-01-13 03:30:00   
1      101003148266  2020-01-01 00:54:48  2020-01-01 00:54:48   
2      101003148268  2020-01-01 01:14:13  2020-01-03 03:30:00   
3      101003148269  2020-01-01 01:19:00  2020-01-03 03:30:00   
4      101003148271  2020-01-01 02:02:00  2020-01-03 03:30:00   
5      101003148272  2020-01-01 02:04:00  2020-01-03 03:30:00   
6      101003148273  2020-01-01 02:11:10  2020-01-03 03:30:00   
7      101003148274  2020-01-01 02:19:58                  NaN   
8      101003148276  2020-01-01 02:56:41  2020-01-03 03:30:00   
9      101003148277  2020-01-01 03:08:52  2020-01-06 03:30:00   
10     101003148278  2020-01-01 03:15:14  2020-01-03 03:3
```

```python
# Ensure the "open_dt" column is treated as a string (in case it isn't already)
data['open_dt'] = data['open_dt'].astype(str)

# Extract the first four characters from the "open_dt" column and create a new column "Year"
data['Year'] = data['open_dt'].str[:4]

# Check the updated DataFrame
print(data[['open_dt', 'Year']].head())

```

Output:
```
open_dt  Year
0  2020-01-01 00:07:00  2020
1  2020-01-01 00:54:48  2020
2  2020-01-01 01:14:13  2020
3  2020-01-01 01:19:00  2020
4  2020-01-01 02:02:00  2020
```

```python
data.to_csv("combined_with_year.csv", index=False)
```

```python

```

Tests for q10 (points: None)
```
>>> assert type(ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))) == dict
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['1_norm'] == 9.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['inf_norm'] == 3.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['2_norm'] > 5.0
```


Tests for q2 (points: None)
```
>>> assert len(roll_die(10)) == 10
>>> assert all((1 <= x <= 6 for x in roll_die(100)))
```


Tests for q3 (points: None)
```
>>> assert answer3() > 0.15 and answer3() < 0.2
>>> assert type(answer3()) == float
```


Tests for q4 (points: None)
```
>>> assert type(ans4()) == float
>>> assert 2 < ans4() < 2.1
```


Tests for q5 (points: None)
```
>>> assert type(ans5() == float)
>>> assert abs(ans5() - 0.49) < 1e-09
>>> assert 0.49 < ans5() < 0.491
```


Tests for q6 (points: None)
```
>>> assert type(ans6()) == float
>>> assert 0.2 < ans6() < 0.21
>>> assert abs(ans6()) - 0.2 < 0.05
```


Tests for q7 (points: None)
```
>>> assert type(ans7()) == float
>>> assert 0.1 < ans7() < 0.11
>>> assert abs(ans7()) - 0.1 < 0.0042
```


Tests for q8 (points: None)
```
>>> assert type(ans8()) == float
>>> assert 0.04 < abs(ans8()) < 0.05
```


Tests for q9 (points: None)
```
>>> assert ans9({1, 2, 3, 4, 5}, {2, 3, 4}) == 0.6
>>> assert 0.4 < ans9({1, 2, 3, 4, 5, 6, 7}, {2, 3, 4}) < 0.5
>>> assert 0.2 < ans9({1, 2, 3, 4, 5, 6, 7, 8, 9}, {2, 3}) < 0.23
```
