```python
!pip install pandas openpyxl
```

Output:
```
Requirement already satisfied: pandas in /opt/anaconda3/envs/scottenvi/lib/python3.12/site-packages (2.2.2)
Collecting openpyxl
  Downloading openpyxl-3.1.5-py2.py3-none-any.whl.metadata (2.5 kB)
Requirement already satisfied: numpy>=1.26.0 in /opt/anaconda3/envs/scottenvi/lib/python3.12/site-packages (from pandas) (1.26.4)
Requirement already satisfied: python-dateutil>=2.8.2 in /opt/anaconda3/envs/scottenvi/lib/python3.12/site-packages (from pandas) (2.9.0.post0)
Requirement already satisfied: pytz>=2020.1 in /opt/anaconda3/envs/scottenvi/lib/python3.12/site-packages (from pandas) (2024.1)
Requirement already satisfied: tzdata>=2022.7 in /opt/anaconda3/envs/scottenvi/lib/python3.12/site-packages (from pandas) (2023.3)
Collecting et-xmlfile (from openpyxl)
  Downloading et_xmlfile-1.1.0-py3-none-any.whl.metadata (1.8 kB)
Requirement already satisfied: six>=1.5 in /opt/anaconda3/envs/scottenvi/lib/python3.12/site-packages (from python-dateutil>=2.8.2->pandas) (1.16.0)
Downloading openp
```

```python
import pandas as pd
import os

# List of directories (folder names contain the year information)
folders = [
    '2018-2019 School Year Reports',  # Replace with the actual folder names
    '2019-2020 School Year Reports'   # Replace with the actual folder names
]

# The sheet name you want to extract from each file
sheet_name = 'Student Addresses'

# Create an empty list to store dataframes
df_list = []

# Function to extract the year from the folder name
def extract_year_from_folder(folder_name):
    # Assumes year is at the beginning of the folder name, e.g., '2018-2019 School Year Reports'
    year = folder_name.split(' ')[0]
    return year

# Function to extract the university name from the file name
def extract_university_from_file(file_name):
    # Remove the file extension (.xlsx) and split by underscore '_'
    university = file_name.split('.')[0].split('_')[0]
    return university

# Loop through all folders
for folder_path in folders:
    # Extract year from the folder name
    year_info = extract_year_from_folder(folder_path)
    
    # List all files in the directory
    files = [file for file in os.listdir(folder_path) if file.endswith('.xlsx')]

    # Loop through all files in each folder and append the specific sheet to the list
    for file in files:
        file_path = os.path.join(folder_path, file)
        
        # Extract university from the file name
        university_info = extract_university_from_file(file)
        
        # Read the specific sheet
        df = pd.read_excel(file_path, sheet_name=sheet_name)
        
        # Check if the DataFrame is not empty
        if not df.empty:
            # Add 'Year' and 'University' columns to the DataFrame
            df['Year'] = year_info
            df['University'] = university_info
            
            # Append the dataframe with the new columns to the list
            df_list.append(df)

# Check if there are any DataFrames to combine
if df_list:
    # Concatenate all dataframes into one
    combined_df = pd.concat(df_list, ignore_index=True)

    # Save the combined dataframe to a new Excel file
    combined_df.to_excel('combined_specific_sheet.xlsx', index=False)

    print("All specific sheets combined successfully into 'combined_specific_sheet.xlsx'.")
else:
    print("No data found to combine.")

```

Output:
```
/opt/anaconda3/envs/scottenvi/lib/python3.12/site-packages/openpyxl/worksheet/_reader.py:329: UserWarning: Unknown extension is not supported and will be removed
  warn(msg)
All specific sheets combined successfully into 'combined_specific_sheet.xlsx'.
```

```python

```

Tests for q10 (points: None)
```
>>> assert type(ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))) == dict
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['1_norm'] == 9.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['inf_norm'] == 3.0
>>> assert ans10(np.array([1, 2, 3]), np.array([4, 5, 6]))['2_norm'] > 5.0
```


Tests for q2 (points: None)
```
>>> assert len(roll_die(10)) == 10
>>> assert all((1 <= x <= 6 for x in roll_die(100)))
```


Tests for q3 (points: None)
```
>>> assert answer3() > 0.15 and answer3() < 0.2
>>> assert type(answer3()) == float
```


Tests for q4 (points: None)
```
>>> assert type(ans4()) == float
>>> assert 2 < ans4() < 2.1
```


Tests for q5 (points: None)
```
>>> assert type(ans5() == float)
>>> assert abs(ans5() - 0.49) < 1e-09
>>> assert 0.49 < ans5() < 0.491
```


Tests for q6 (points: None)
```
>>> assert type(ans6()) == float
>>> assert 0.2 < ans6() < 0.21
>>> assert abs(ans6()) - 0.2 < 0.05
```


Tests for q7 (points: None)
```
>>> assert type(ans7()) == float
>>> assert 0.1 < ans7() < 0.11
>>> assert abs(ans7()) - 0.1 < 0.0042
```


Tests for q8 (points: None)
```
>>> assert type(ans8()) == float
>>> assert 0.04 < abs(ans8()) < 0.05
```


Tests for q9 (points: None)
```
>>> assert ans9({1, 2, 3, 4, 5}, {2, 3, 4}) == 0.6
>>> assert 0.4 < ans9({1, 2, 3, 4, 5, 6, 7}, {2, 3, 4}) < 0.5
>>> assert 0.2 < ans9({1, 2, 3, 4, 5, 6, 7, 8, 9}, {2, 3}) < 0.23
```
